**Question 1**

(This is Problem 4.4 from Gregory.) A bottle contains three green balls and three red balls.
The bottle is first shaken to mix up the balls.

(a) (2 pts) Calculate analytically the probability that blindfolded, you will pick a red ball
on the third pick, if you learn that at least one red ball was picked on the first two picks.

(b) CODING: (3 pts) Write a Python function that simulates this an instance of drawing
three balls from the bottle.

(c) CODING: (3 pts) Using the function from b), empirically calculate the probability
from part a).

<!-- grade_id: cell-a8606aa986a5e617, points: 2, grade: true, solution: true -->
Quesiton 1a:

YOUR ANSWER HERE:
Let A be the event that the third ball picked is red.
Let B be the event that at least one red ball was picked in the first two picks.
We need to calculate P(A|B)= P(A and B)/ P(B)
P(B) = 1- P(!B)
P(!B) is getting the first two as green = (1/2)*(2/5) = 1/5
P(B) = 4/5
P(A and B) = probability of first 3 as red or first and third red or second and third red= (1/2)*(2/5)*(1/4) + (1/2)*(3/5)*(1/2) +(1/2)*(3/5)*(1/2) =(7/20)

So P(A|B) = (7/20)/(4/5) = 7/16 = 0.4375
        

<!-- grade_id: cell-a29bdea1673f34c6, points: 0, grade: false, solution: true -->
```python
# Question 1b

import numpy as np
import random
def simulate_draw3():
    """Simulate drawing 3 balls from a bottle containing 3 green balls and 3 red balls at random
    
    Returns: an array of three strings, each of which is either 'G' or 'R',
             representing green and red respectively.
    """
    Red =0
    Green = 0
    output = []
    for i in range(3):
        num = random.random()
        if num < (3-Red)/(6-(Red+Green)):
            output.append('R')
            Red+=1
        else:
            output.append('G')
            Green+=1
    return output

```

<!-- grade_id: cell-e2d5c2ede65447c8, points: 3, grade: true, solution: false -->
```python
draw = simulate_draw3()

# If you run this multiple times, this should (usually) change each time.
print(f'The three draws are: {draw}')

# Some sanity checks:

# Length should be 3
assert len(draw) == 3

# Each one should be either 'G' or 'R'
d1, d2, d3 = draw
assert d1 in ['G', 'R']
assert d2 in ['G', 'R']
assert d3 in ['G', 'R']

```

Output:
```
The three draws are: ['G', 'R', 'R']
```

<!-- grade_id: cell-fbb32c8fb25debed, points: 0, grade: false, solution: true -->
```python
# Question 1c

def estimate_prob():
    """Estimate the probability of drawing a red ball 3rd given that at least one of the first two
    draws were red, by runing simulate_draw3 many times.
    
    Returns 
    """
    validstates = 0
    Redonthird=0
    for i in range(1000000):
        state = simulate_draw3()
        if (state[0] == 'R' or state[1] == 'R'):
            validstates+=1
            if state[2]== 'R':
                Redonthird +=1
    return Redonthird/validstates
```

<!-- grade_id: cell-8832a8df51899375, points: 3, grade: true, solution: false -->
```python
prob = estimate_prob()
print(f'The estimated probability is {prob}')

# You should compare this to what you got in Question 1a.  It should be equal to 2 decimal places.

```

Output:
```
The estimated probability is 0.4368871966920496
```

**Question 2**

(6 pts; this is Problem 4.7 from Gregory.) In a particular water sample, ten bacteria are
found, of which three are of type $A$. Calculate analytically the probability of obtaining six
type $A$ bacteria, in a second independent water sample containing 12 bacteria in total.

<!-- grade_id: cell-446be4841a80674a, points: 2, grade: true, solution: true -->
Hint: You may find the following definite integral to be helpful:

$$
\int_0^1  x^k (1-x)^{n-k} dx = \frac{1}{(n+1)\binom{n}{k}}
$$

YOUR ANSWER HERE
Let A be the probability of getting 6 bacteria out of 12
and let D be our data and P be our parameter

Our equation is 
$$
P(A|D,P,I) = \frac{P(A|P,I)\times P(D|A,P,I)}{P(D|P,I)} = \frac{likelyhood\times Prior}{normalization constant}
$$
First we will assign our prior as 
$$
prior = P(D|A,P,I) = \binom{10}{3}\times p^3 \times (1-p)^7
$$
and our liklehood function as
$$
likelyhood = P(A|P,I)= \binom{12}{6}\times p^6 \times (1-p)^6
$$
and our normalization factor is 
$$
normalization factor = P(D|I) = \binom{10}{3}\times p^3 \times (1-p)^7
$$
Now lets plug all of this in the equation
$$
P(A|D,P,I) = \frac{
\int_0^1 P(A|P,I)\times P(D|A,P,I)}{\int_0^1P(D|P,I)} = \binom{12}{6}\times\binom{10}{3} \times\frac{\int_0^1 p^9 (1-p)^13 dp}{\binom{10}{3}\int_0^1 p^3 (1-p)^7}
$$
Which will evaluate to
$$
P(A|D,P,I) = \frac{\binom{12}{6}\times\binom{10}{3}\times 11}{23\times\binom{22}{9}} = \frac{792}{7429} = 0.1066 = 10.7\%
$$

**Question 3**
    
This problem demonstrates empirically several properties of counting statistics, which obeys the 
Poisson distribution. We will approximate a Poisson process by generating random values between
0 and 10, but only looking at the ones between 0 and 1.  Technically, this is a multinomial distribution,
but for large $N$ and at least moderately large number of bins, it is a good approximation to
a Poisson process.

(a) (3 pts) Generate $N = 10^6$ uniform random deviates between 0 and 10.
    Count the number of deviates that fall in each of 100 equal-sized bins between 0 and 1 
    (0.00-0.01, 0.01-0.02, etc).  Calculate the mean $\mu$ and variance $\sigma^2$ of the counts per bin.
    Make sure that the mean is almost exactly what you expect: $\bar{n} \approx N/100 = 1000$.

(b) (3 pts) Repeat part (a) for 10 and 1000 equally-sized bins.  This empirically shows that $\mu = \sigma^2 = \bar{n}$, 
    which is an important property of the Poisson distribution.
    
(c) **EXTRA CREDIT** (2 pts): Now generate $M = 100$ realizations of part (a).
    Make sure to not use the same random seed each time.  You should now have $M = 100$ values of both the mean and 
    variance.  Let's call these $\mu_i$ and $\sigma_i^2$, respectively, where $i = 1, 2, ..., M$.
    Now calculate the mean and variance of $\mu_i$ and $\sigma_i^2$.  This demonstrates that the general property 
    that the variance (2nd moment) is much more difficult to measure accurately than the mean (1st moment).

<!-- grade_id: cell-30668e3a915c451c, points: 0, grade: false, solution: true -->
```python
# Question 3a,b

def make_counts(ntot, nbins):
    """Make a random array of ntot uniform random deviates between 0 and 10, 
    and bin the values between 0 and 1 into nbins bins.
    
    Returns the number of values in each bin as an array of lenth nbins
    """
    # Hints:
    # 1. Use either np.random.random or np.random.uniform to make the full array.
    # 2. There are a number of ways to bin these.  Probably np.histogram is easiest.
    #    Read the docs carefully, since you'll need to set some non-default parameters.
    
    random_deviates = np.random.uniform(0, 10, ntot)
    bin_edges = np.arange(0, 1.01, 1/nbins)
    counts, _ = np.histogram(random_deviates, bins=bin_edges)
    
    return counts
    
def calculate_mean_var(x):
    """Return the mean and variance of the values in the given array (as a tuple).
    """
    # You've done this twice already, so this should be a trivial one-liner at this point.
    # Feel free to use numpy functions for this.
    
    mean_counts = np.mean(x)
    variance_counts = np.var(x)
    return (mean_counts,variance_counts)

#part(a)
N = 10**6
num_bins_a = 100

counts_per_bin_a = make_counts(N, num_bins_a)
mean_a, variance_a = calculate_mean_var(counts_per_bin_a)

expected_mean_a = N / num_bins_a
# Output
print("Part (a):")
print(f"Mean counts per bin (μ): {mean_a}")
print(f"Variance of counts per bin (σ²): {variance_a}")
#expected always greater by multiplication of 10 so I will divide by 10 :)
print(f"Expected mean: {expected_mean_a/10}\n")

#part(b)
N = 10**6
bin_sizes_b = [10, 1000]

results_b = {}
for num_bins in bin_sizes_b:
    counts_per_bin = make_counts(N, num_bins)
    mean, variance = calculate_mean_var(counts_per_bin)
    expected_mean = N / num_bins  # Expected mean for Poisson distribution
    results_b[num_bins] = (mean, variance, expected_mean)
# Output
print("Part (b):")
for num_bins, (mean, variance, expected_mean) in results_b.items():
    print(f"Number of bins: {num_bins}")
    print(f"Mean counts per bin (μ): {mean}")
    print(f"Variance of counts per bin (σ²): {variance}")
    #expected always greater by multiplication of 10 so I will divide by 10 :)
    print(f"Expected mean: {expected_mean/10}\n")
```

Output:
```
Part (a):
Mean counts per bin (μ): 1001.99
Variance of counts per bin (σ²): 927.4898999999998
Expected mean: 1000.0

Part (b):
Number of bins: 10
Mean counts per bin (μ): 9978.7
Variance of counts per bin (σ²): 12861.61
Expected mean: 10000.0

Number of bins: 1000
Mean counts per bin (μ): 99.46382556987116
Variance of counts per bin (σ²): 102.58962302606571
Expected mean: 100.0
```

<!-- grade_id: cell-78364ea94b9eec92, points: 6, grade: true, solution: false -->
```python
N = 1_000_000 # The underscores are ignored. For large integers, they are allowed to make it easier to read.

# 4a:
nbins = 100

x = make_counts(N, nbins)
mean, var = calculate_mean_var(x)

print(f"Mean of {nbins} values = {mean:.2f}")
print(f"Expected mean = {N/10/nbins:.2f}")

print(f"Variance of {nbins} values = {var:.2f}")
print(f"Expected variance = {N/10/nbins:.2f}")

# 4b:
nbins = 10
x = make_counts(N, nbins)
mean, var = calculate_mean_var(x)

print()
print(f"Mean of {nbins} values = {mean:.2f}")
print(f"Expected mean = {N/10/nbins:.2f}")

print(f"Variance of {nbins} values = {var:.2f}")
print(f"Expected variance = {N/10/nbins:.2f}")

nbins = 1000
x = make_counts(N, nbins)
mean, var = calculate_mean_var(x)

print()
print(f"Mean of {nbins} values = {mean:.2f}")
print(f"Expected mean = {N/10/nbins:.2f}")

print(f"Variance of {nbins} values = {var:.2f}")
print(f"Expected variance = {N/10/nbins:.2f}")



```

Output:
```
Mean of 100 values = 997.89
Expected mean = 1000.00
Variance of 100 values = 1254.66
Expected variance = 1000.00

Mean of 10 values = 9950.90
Expected mean = 10000.00
Variance of 10 values = 10318.69
Expected variance = 10000.00

Mean of 1000 values = 100.23
Expected mean = 100.00
Variance of 1000 values = 105.95
Expected variance = 100.00
```

<!-- grade_id: cell-6b7208ff159c9567, points: 0, grade: false, solution: true -->
```python
# Problem 3c:

def run_simulation(M, N, nbins):
    """Run M realizations of the calculation done for Problem 4a
    
    Returns two lists (mu, sigsq), each of length M.
    """
    mu = []
    sigsq = []
    for i in range(M):
        np.random.seed(i)
        counts_per_bin = make_counts(N, nbins)
        mean, variance = calculate_mean_var(counts_per_bin)
        mu.append(mean)
        sigsq.append(variance)
    return mu, sigsq
M = 100
N = 10**6
num_bins = 100
mu_values, sigma_squared_values = run_simulation(M, N, num_bins)
mean_mu = np.mean(mu_values)
var_mu = np.var(mu_values)
mean_sigma_squared = np.mean(sigma_squared_values)
var_sigma_squared = np.var(sigma_squared_values)

# Output
print("Results from 100 realizations:")
print(f"Mean of μ: {mean_mu}")
print(f"Variance of μ: {var_mu}")
print(f"Mean of σ²: {mean_sigma_squared}")
print(f"Variance of σ²: {var_sigma_squared}")
```

Output:
```
Results from 100 realizations:
Mean of μ: 1000.3209000000003
Variance of μ: 10.30862018999999
Mean of σ²: 993.8991029999999
Variance of σ²: 27226.783107438292
```

<!-- grade_id: cell-ffbd9ac6df75e7a0, points: 2, grade: true, solution: false -->
```python
mu, sigsq = run_simulation(100, 1_000_000, 100)

mu_mean, mu_var = calculate_mean_var(mu)
sigsq_mean, sigsq_var = calculate_mean_var(sigsq)

print(f"The mean and variance of mu = {mu_mean:.2f}, {mu_var:.2f}")
print(f"The mean and variance of sigsq = {sigsq_mean:.2f}, {sigsq_var:.2f}")

print(f"I.e. the mean values fall mostly in the range {mu_mean:.2f} +- {np.sqrt(mu_var):.2f}")
print(f"and the sigsq values fall mostly in the range {sigsq_mean:.2f} +- {np.sqrt(sigsq_var):.2f}")

```

Output:
```
The mean and variance of mu = 1000.32, 10.31
The mean and variance of sigsq = 993.90, 27226.78
I.e. the mean values fall mostly in the range 1000.32 +- 3.21
and the sigsq values fall mostly in the range 993.90 +- 165.01
```